<a href="https://colab.research.google.com/github/Chris-Beaudoin/Crypto-Dashboard/blob/main/ps3.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

# Problem Set 3
#### CMSI 1010: Computer Programming and Lab

# Important Dates

 - Released: 2022-09-29 at 12:00 am [PT]
 - Deadline: 2022-10-06 at 11:59 pm [PT]

# Question 1: Reading Notes

All readings are from the [_Think Python_ textbook](http://greenteapress.com/thinkpython2/thinkpython2.pdf):
* Supporting content for this pset: Chapters 7–9
* Upcoming content: Chapters 11–13, 15–17

Take notes from your readings and submit them as part of this pset.

A common kind of reassignment is an update, where the new value of the variable depends
on the old

Before you can update a variable, you have to initialize it, usually with a simple assignment:

Updating a variable by adding 1 is called an increment; subtracting 1 is called a decrement.

Loop that does not stop is called an infinite loop

the break statment ends the passing through of the loop and continue passes to the next variable

Loops are often used in programs that compute numerical results by starting with an approximate answer and iteratively improving it

Newton’s method is an example of an algorithm:

reassignment: Assigning a new value to a variable that already exists.
update: An assignment where the new value of the variable depends on the old.
initialization: An assignment that gives an initial value to a variable that will be updated.
increment: An update that increases the value of a variable (often by one).
decrement: An update that decreases the value of a variable.
iteration: Repeated execution of a set of statements using either a recursive function call
or a loop.
infinite loop: A loop in which the terminating condition is never satisfied.
algorithm: A general process for solving a category of problems

A string is a sequence, which means it is
an ordered collection of other values. In this chapter you’ll see how to access the characters
that make up a string, and you’ll learn about some of the methods strings provide

The expression in brackets is called an index. The index indicates which character in the
sequence you want (hence the name).

len is a built-in function that returns the number of characters in a string:

Or you can use negative indices, which count backward from the end of the string. The
expression fruit[-1] yields the last letter, fruit[-2] yields the second to last, and so on

This pattern of processing is called a traversal. One way to write a traversal is with a
while loop:

A segment of a string is called a slice

If the first index is greater than or equal to the second the result is an empty string, represented by two quotation marks

The reason for the error is that strings are immutable, which means you can’t change an
existing string. The best you can do is create a new string that is a variation on the original:

This pattern of computation—traversing a sequence and returning when we find what we
are looking for—is called a search.

This program demonstrates another pattern of computation called a counter. The variable
count is initialized to 0 and then incremented each time an a is found. When the loop exits,
count contains the result—the total number of a’s.

A method call is called an invocation; in this case, we would say that we are invoking
upper on word.

This is an example of an optional argument; find can also take a third argument, the index
where it should stop:

The word in is a boolean operator that takes two strings and returns True if the first appears as a substring in the second:

The relational operators work on strings. To see if two strings are equal:

object: Something a variable can refer to. For now, you can use “object” and “value”
interchangeably.
sequence: An ordered collection of values where each value is identified by an integer
index.
item: One of the values in a sequence.
index: An integer value used to select an item in a sequence, such as a character in a string.
In Python indices start from 0.
slice: A part of a string specified by a range of indices.
empty string: A string with no characters and length 0, represented by two quotation
marks.
immutable: The property of a sequence whose items cannot be changed.
traverse: To iterate through the items in a sequence, performing a similar operation on
each.
search: A pattern of traversal that stops when it finds what it is looking for.
counter: A variable used to count something, usually initialized to zero and then incremented.
invocation: A statement that calls a method.
optional argument: A function or method argument that is not required.

There are lots of word
lists available on the Web, but the one most suitable for our purpose is one of the word lists
collected and contributed to the public domain by Grady Ward as part of the Moby lexicon project

This file is in plain text, so you can open it with a text editor, but you can also read it from
Python. The built-in function open takes the name of the file as a parameter and returns a
file object

The for loop traverses the characters in word. If we find the letter “e”, we can immediately
return False; otherwise we have to go to the next letter. If we exit the loop normally, that
means we didn’t find an “e”, so we return True.


For is_abecedarian we have to compare adjacent letters, which is a little tricky with a for
loop:
def is_abecedarian(word):
previous = word[0]
for c in word:
if c < previous:
return False
previous = c
return True
An alternative is to use recursion:
def is_abecedarian(word):
if len(word) <= 1:
return True
if word[0] > word[1]:
return False
return is_abecedarian(word[1:])
Another option is to use a while loop:
def is_abecedarian(word):
i = 0
while i < len(word)-1:
if word[i+1] < word[i]:
return False
i = i+1
return True

t. To
convince yourself that the loop ends correctly, consider an example like 'flossy'. The
length of the word is 6, so the last time the loop runs is when i is 4, which is the index of
the second-to-last character. On the last iteration, it compares the second-to-last character
to the last, which is what we want.

The empty string is an example of a special case, which is one of the non-obvious
cases where errors often lurk.

file object: A value that represents an open file.
reduction to a previously solved problem: A way of solving a problem by expressing it
as an instance of a previously solved problem.
special case: A test case that is atypical or non-obvious (and less likely to be handled correctly).



# Question 2: Understanding loops

### Part A
 For these fragments of code, write the output **_without_** running them.

```python
    num1 = 10
    while num1 > 3:
        print(num1)
        num1 = num1 - 1
```

 #### Output:
    ```
    10
    9 
    8 
    7 
    6 
    5 
    4
    ```

```python
    num2 = 10
    while True:
        if num2 < 7:
            # break makes us quit the loop
            break
        print(num2)
        num2 = num2 - 1
```

 #### Output:
    ```
    10
    9
    8
    7
    ```

```python
    divisor = 2
    for i in range(0, 10, 2):
        print(i / divisor)
```

 #### Output:
    ```
    <!-- 1/0
    2/10
    2/3??? -->

    1.0
    2.0
    3.0
    4.0
    ```

### Part B

These programs use `for` loops to iterate through strings.

```python
    my_string = 'stressed'
    for letter in my_string:
        print(letter)
```

#### Output:
    ```
    s
    t
    r
    e
    s
    s
    e
    d
    ```

```python
    my_string = 'stressed'
    res = ''
    for letter in my_string:
        res = letter + res
    print(res)
```

#### Output:
    ```
    desserts
    ```

Describe, in natural language, what’s happening in this loop:

(edit this cell with your answer)

# Question 3 : Exponentials

Imagine that the `**` operator for numbers has now been removed from the Python programming language. Write a program that prompts the user for a base and an exponent and uses a `for` loop to calculate `base ** exp`. _**Note:** You should handle the case where the user does not type a positive integer._

```python
def exponential(base, exp):
    # TODO: add your code here for calculating the final result
    result = 1
    iniExp = exp
    exp = abs(exp)
    if exp == 0:
      return 1
    for f in range(exp):
        result *= base
    if iniExp < 0:
      return (base / result)/2
    return result

    # for num in range(exp - 1): 
    #   result = (result
    #             * base)
    # print(result)
    # return result


exponential(4, 10)
```

Output:
```
1048576
```

```python
# TODO: ask for user input, then call the exponential function
base = eval(input("What would you like to be your base value?"))
```

Output:
```
What would you like to be your base value?11
```

The cells below are autograder cells; don't edit those! But you can run them (using the Run button) to test your own code! If they don't spit out any error messages, then your code works!

```python
assert exponential(2,4) == 2**4
```

```python
assert exponential(3,0) == 3**0
```

```python
assert exponential(-2,3) == (-2)**3
```

```python
assert exponential(2,-2) == 2**(-2)
```

```python
import random

num1 = random.randint(0,1000)
num2 = random.randint(0,100)

assert exponential(num1, num2) == num1**num2
```

# Question 4 : Functions and more functions

### Part A

Write a function:

 - named `list_of_positive_indices`
 - takes 1 argument, a list
 - returns 1 value, a list
 
In this function, you should loop through each element of the argument list. If the element is *positive*, then you should append the *index number of the positive element* to a new list. If not, skip that index (i.e. do nothing). Once the loop is finished, return the new list.

For example, `list_of_positive_indices([1, -1, 0, 3])` should return `[0, 3]`. No imports!

*HINT*: Instead of doing a simple `for element in list:` loop, try using the `range()` function in the loop header.

```python
def list_of_positive_indices(in_list):
    new_list = []
    for number in in_list:
      if number > 0:
        new_list.append(number)
      else:
        continue
      print(new_list)

# print(new_list)

list_of_positive_indices([1, -1, 0, 3])

```

Output:
```
[1]
[1, 3]
```

Run the autograder cells below to test your code!

<!-- grade_id: q1c_test1, points: 5, grade: true, solution: false -->
```python
import numpy as np
np.random.seed(986)

list1 = np.random.randint(-100, 100, 10)
indices1 = np.arange(list1.shape[0])[list1 > 0].tolist()
pred1 = list_of_positive_indices(list1.tolist())
assert set(pred1) == set(indices1)
```

Output:
```
[51]
[51, 43]
[51, 43, 22]
[51, 43, 22, 89]
[51, 43, 22, 89, 2]
TypeError                                 Traceback (most recent call last)
<ipython-input-15-8f786458eab7> in <module>
      5 indices1 = np.arange(list1.shape[0])[list1 > 0].tolist()
      6 pred1 = list_of_positive_indices(list1.tolist())
----> 7 assert set(pred1) == set(indices1)

TypeError: 'NoneType' object is not iterable
```

<!-- grade_id: q1c_test2, points: 5, grade: true, solution: false -->
```python
np.random.seed(578752)

list2 = np.random.randint(-100, 100, 100)
indices2 = np.arange(list2.shape[0])[list2 > 0].tolist()
pred2 = list_of_positive_indices(list2.tolist())
assert set(pred2) == set(indices2)
```

Output:
```
[65]
[65, 57]
[65, 57, 19]
[65, 57, 19, 57]
[65, 57, 19, 57, 78]
[65, 57, 19, 57, 78, 12]
[65, 57, 19, 57, 78, 12, 20]
[65, 57, 19, 57, 78, 12, 20, 26]
[65, 57, 19, 57, 78, 12, 20, 26, 74]
[65, 57, 19, 57, 78, 12, 20, 26, 74, 24]
[65, 57, 19, 57, 78, 12, 20, 26, 74, 24, 49]
[65, 57, 19, 57, 78, 12, 20, 26, 74, 24, 49, 74]
[65, 57, 19, 57, 78, 12, 20, 26, 74, 24, 49, 74, 34]
[65, 57, 19, 57, 78, 12, 20, 26, 74, 24, 49, 74, 34, 72]
[65, 57, 19, 57, 78, 12, 20, 26, 74, 24, 49, 74, 34, 72, 48]
[65, 57, 19, 57, 78, 12, 20, 26, 74, 24, 49, 74, 34, 72, 48, 46]
[65, 57, 19, 57, 78, 12, 20, 26, 74, 24, 49, 74, 34, 72, 48, 46, 26]
[65, 57, 19, 57, 78, 12, 20, 26, 74, 24, 49, 74, 34, 72, 48, 46, 26, 1]
[65, 57, 19, 57, 78, 12, 20, 26, 74, 24, 49, 74, 34, 72, 48, 46, 26, 1, 69]
[65, 57, 19, 57, 78, 12, 20, 26, 74, 24, 49, 74, 34, 72, 48, 46, 26, 1, 69, 16]
[65, 57, 19, 57, 78, 12, 20, 26, 74, 24, 49, 74, 34, 72, 48, 46, 26, 1, 69, 16, 84]
[65, 57, 19, 57, 78, 12, 20, 26, 74, 24, 49, 74, 34, 72, 48
```

### Part B

Write a function that tests if a number is even.

*HINT*: Recall the modulo operator `%`. This performs division on two numbers, but instead of returning their quotient, it returns the *remainder* of the division. For example, `7 % 3` returns 1, since the quotient is 2 with a remainder of 1.

Your function should:

 - be named `is_even`
 - take 1 argument, a number
 - return `True` if the argument is even, `False` otherwise
 
Make sure you're returning a *boolean*!

<!-- grade_id: q2a, points: 0, grade: false, solution: true -->
```python
def is_even(num):
    if num % 2 == 0:
      return True
    else:
      return False
```

<!-- grade_id: q2a_test1, points: 2.5, grade: true, solution: false -->
```python
assert is_even(3) == False
```

<!-- grade_id: q2a_test2, points: 2.5, grade: true, solution: false -->
```python
assert is_even(123456789) == False
```

<!-- grade_id: q2a_test3, points: 2.5, grade: true, solution: false -->
```python
assert is_even(42) == True
```

<!-- grade_id: q2a_test4, points: 2.5, grade: true, solution: false -->
```python
assert is_even(42.1) == False
```

### Part C

How could you test if a number is *odd*, using your `is_even` function from Part B?

Write your answer, either using 1 line of code or explaining in only 1 sentence, in the box below (double click on the statement below to activate the box).

<!-- grade_id: q2b, points: 5, grade: true, solution: true -->
you could test if it is odd by getting the remainder and checking if it is equal to one.

### Part D *(Optional)*

*This part is optional*. Write a function that finds the *second-smallest* number in a list of numbers. Your function should:

 - be named `second_smallest`
 - take 1 argument, a list of numbers
 - return 1 number, the second-smallest number in the list
 
Take note of that last point. There's a possibility, after all, that the list your function receives as input will have 1 or even 0 elements in it. How, then, do you return the second-smallest element of a 0-element list? Well, simply, you can't. So instead, you'll return False.

For example, `second_smallest([1, 2, 3])` should return `2`. Another example `second_smallest([10, 100, 1000, 10000])` should return 100. And `second_smallest([1])` should returb `False`.

You **cannot** use any built-in sorting functions! Only `range()` and `len()`.

<!-- grade_id: q2c, points: 0, grade: false, solution: true -->
```python
def second_smallest(in_list):
    pass
```

<!-- grade_id: q2c_test1, points: 5, grade: true, solution: false -->
```python
import numpy as np
np.random.seed(89547)

l1 = np.random.randint(-100, 100, 100)
a1 = np.sort(l1)[1]
np.testing.assert_allclose(a1, second_smallest(l1.tolist()))
```

<!-- grade_id: q2c_test2, points: 5, grade: true, solution: false -->
```python
np.random.seed(485)

l2 = np.random.randint(0, 1000, 100)
a2 = np.sort(l2)[1]
np.testing.assert_allclose(a2, second_smallest(l2.tolist()))
```

<!-- grade_id: q2c_test3, points: 2.5, grade: true, solution: false -->
```python
assert second_smallest([]) == False
```

<!-- grade_id: q2c_test4, points: 2.5, grade: true, solution: false -->
```python
assert second_smallest([100]) == False
```

# Question 5: Greetings

### Part A

Print a greeting saying “Hello, Harry” and “Goodbye, Potter” using only the variables provided below. Complete instructions below.

```python
# Write a series of print statements to recreate a 
# greeting box and a goodbye box for Harry Potter.

# *******************
# *                 *
# * Hello, Harry!!! *
# *                 *
# *******************
#
# ***********************
# *                     *
# * Goodbye, Potter!!!  *
# *                     *
# ***********************


name = 'Harry Potter'
greeting = 'Hello, world!'
goodbye = 'Goodbye, all!'
space = ' '
star = '*'

name = name.split(' ')
greeting = greeting.split(' ')
goodbye = 'Goodbye, Potter!!!'
goodbye = goodbye.split(' ')
name = greeting[0] + name[0]
name = name.replace(',', ' ')
greeting = "!!!"
name = str(name) + str(greeting)
print(name)
# goodbye = str(goodbye).replace('[', '')
# goodbye = str(goodbye).replace(']', '')
goodbye = goodbye[0] + goodbye[1]
print(goodbye)
# You may only use the variables defined above. 
# For example, if you wanted to print
# a "T" in stars, you would do:

print (star * 20)
print (star + space * 18 + star)
print (star + space*2 + name + space*2 + star)
print (star + space * 18 + star)
print (star * 20)

print (star * 20)
print (star + space * 18 + star)
print (star + space + goodbye + star)
print (star + space * 18 + star)
print (star * 20)

# print (star * 3)
# print (space + star + space)
# print (space + star + space)

# The answer below would not be acceptable:
# print ('***')
# print (' * ')
# print (' * ')
```

Output:
```
Hello Harry!!!
Goodbye,Potter!!!
********************
*                  *
*  Hello Harry!!!  *
*                  *
********************
********************
*                  *
* Goodbye,Potter!!!*
*                  *
********************
```

```python

```

### Part B
 
 **Reflection:** The output of `greetings` looks very similar to the one from PS0. However, the code looks different. How would you compare the two versions? What advantages/disadvantages do you see in the two approaches?

The new one is far easier to edit and add to especially when editing multiple at once, wheras the other way around would need to be edited each in specific.

# Question 6: Acronym

This exercise will allow you to practice with the methods for strings and lists. An acronym is a word formed by taking the first letters of the words in a phrase and making a word from them. For example, RAM is an acronym for “random access memory.” In this exercise, you are going to write a program that should allow the user to type in a phrase, and the computer should then output the acronym for that phrase. An example of how your program should work is as follows:

```
    Enter a phrase: loyola marymount university
    Your acronym is: LMU
```

**Note:** The acronym should be all uppercase, even if the words in the phrase are not capitalized.

```python
# TODO: Ask the user to input a phrase.
# TODO: Print out an all-caps acronym for the phrase.

# phrase = eval(input('Write your phrase seperated by spaces'))

acronym = ""
phrase = input("Enter a phrase: ")
new_phrase =(phrase.replace('of','')).replace('and','').split()
for word in new_phrase:
    acronym = acronym + word[0].upper()
print('Your acronym is: ', acronym)
```

Output:
```
Enter a phrase: PIhao aoihihs
Your acronym is:  PA
```